# EC Notebook for Lecture 19: Bayes Rule and `if` Statements

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30 am on **Monday, Oct. 14** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

# Question 1: Bayes Rule

Suppose a factory run by Wade and Karle has machines I and II that produce “Waflan” phones.
The factory’s record show that:

- Machines I and II produce, respectively, 1% and 2% defective Waflans.
- Out of the total production, machines I and II produce, respectively, 45% and 55% of all Waflans.


## Question 1.1

You are given a random phone (you don't know what machine made it).  What is the probability that you has picked up a **defective** phone?


<!-- grade_id: cell-9578fbde59fb455a, points: 0, grade: false, solution: true -->
```python
P_defective = (0.45 + 1.1)/ 100
P_defective



```

Output:
```
0.0155
```

## Question 1.2

If you picked a defective phone, what is the probability that it was produced by Machine I?


<!-- grade_id: cell-8f102827bea7eb1d, points: 0, grade: false, solution: true -->
```python
P_defective_from_M1 = 0.45 / 1.55
P_defective_from_M1


```

Output:
```
0.2903225806451613
```

<!-- grade_id: cell-ddf5fc5fca795b79, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Question 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(round(P_defective,4) == 0.0155), "The probability of being defective is not correct"
assert(round(P_defective_from_M1,2) == 0.29), "The probability of defective being produced by machine 1 is not correct"


## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

Output:
```
🎉 All tests passed! 🎉
```

# Question 2

In this question we will try to use **if-statement** in our simulation.

Suppose we roll a fair, 6-sided dice.

- If we roll a 1, 2, or 3, we roll a second die that is also six-sided.
- If we roll a 4, 5, or 6, we roll a second die that is **12-sided**.

Write a code to run this simulation 10,000 times. Store your output in the DataFrame named `df` in columns `first` for the first die and `second` for the second die.

<!-- grade_id: cell-3851bc3799a6c2ac, points: 0, grade: false, solution: true -->
```python
import random
import pandas as pd
data = []

for i in range(10000):
    roll1 = random.randint(1,6)
    if roll1 <=3:
        roll2 = random.randint(1,6)
    else:
        roll2 = random.randint(1,12)
        
    d = {'first':roll1,'second':roll2}
    data.append(d)
df = pd.DataFrame(data)





df.sample(5)
```

Output:
```
first  second
4399      6       1
5488      5       9
7803      1       1
3130      6       7
2020      1       2
```

<!-- grade_id: cell-e5e0fbae87fcf069, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Question 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(len(df)==10000), "The simulation must run 10,000 times."

assert( 'first' in df ), "You must have a column labeled `first`."
assert( 'second' in df ), "You must have a column labeled `second`."

assert( len( df[ (df['first'] <= 3) & (df['second'] > 6) ] ) == 0), "If the first die rolls a 3 or less, the second die should never roll above a 6."
assert( len( df[ (df['first'] >= 4) & (df['second'] > 6) ] ) != 0), "If the first die rolls a 4 or more, the second die must be a 12-sided die."


## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

Output:
```
🎉 All tests passed! 🎉
```

## Submit Your Work!

Make sure to **Save and Checkpoint** your notebook, exit Jupyter, and submit your work! :)