## EC Notebook for Lecture 35: Hypothesis Testing & Z Tests in Python

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30 am on **Friday, Nov. 22** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## Hypothesis Testing

In today's lecture, we go over the one sample hypothesis testing and study the two sample test. We know how to find the p-value in one sample hypothesis testing and Z-score in two sample hypothesis testing.

Now, we will still use the `beer` dataset to calculate the p-value or Z-score for hypothesis test. 

In the `beer` dataset, we have the following continous variables: `PercentAlcohol`, `Calories` 
`Carbohydrates`. 

```python
import pandas as pd
df = pd.read_csv('beer-dataset.csv')
```

## Puzzle 1 

Mike is a investigor focusing on the beer market, based on his investigation, he has the true information about the true average amount of `Calories` over all beers is $150$ `Calories` ($\mu = 150$). However, based on the information in the `beer` dataset, the `Calories` could be larger or smaller than $150$. Therefore, we are interested in if there is significantly different of `Calories` between Mike's investigation and `beer` dataset. 

## Puzzle 1.1
Try to calculate the **p-value** by using the definiton of Z-score and P-value. Hint: This is a two-tailed alternative hypothesis, i.e. $\mu_1 \neq \mu_2$

<!-- grade_id: cell-aa6c8c81621207f5, points: 0, grade: false, solution: true -->
```python
import pandas as pd
from scipy.stats import norm
df = pd.read_csv('beer-dataset.csv')

mean = df['Calories'].mean()
se = df['Calories'].std() / (len(df)**0.5)
z_score = (mean - 150) / se
z_score = round(z_score,14)
p_value1 = norm.cdf(z_score)*2
round(p_value1,17)
```

Output:
```
0.00537443931276439
```

## Puzzle 1.2
Try to use the Python library **ztest** to calculate the **p-value**. Hint: This is a two-tailed alternative hypothesis, i.e. $\mu_1 \neq \mu_2$.

<!-- grade_id: cell-114fd8384a4132b8, points: 0, grade: false, solution: true -->
```python
from statsmodels.stats.weightstats import ztest
p_value2 = ztest(df['Calories'],value = 150)[1]


```

<!-- grade_id: cell-a46cec17f40d47f5, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Part Puzzle1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(-2.78372 <= round(z_score,5) <= -2.78365), "z-score is not correct"
assert(0.00537443931276436 <= round(p_value1,17) <= 0.00537443931276440), "p_value1 is not correct"
assert(0.00537443931276443 <= round(p_value2,17) <=  0.00537443931276447), "p_value2 is not correct."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()

```

Output:
```
🎉 All tests passed! 🎉
```

## Puzzle 2

Right now, we have answered the question that whether there is significant difference of average of `Calories` between Mike's investigation and information from the `beer` dataset.

Now, Mike is also curious about is there any significant difference of average of `Calories` between the first **42** beers and last **43** beers, we define them as `df1_Calories` and `df2_Calories` in the following:

```python
df1_Calories = df["Calories"][0:42]
df2_Calories = df["Calories"][42:85]
```

## Puzzle 2.1

Try to calculate the **Z-score** by using the definiton of Z-score. Hint: This is a two-sample hypothesis test with null hypothesis $\bar X_1 \neq \bar X_2$.

<!-- grade_id: cell-ed5b33e8e161aa37, points: 0, grade: false, solution: true -->
```python
x1 = df1_Calories.mean()
x2 = df2_Calories.mean()
se1 = df1_Calories.std() / (len(df1_Calories)**0.5)
se2 = df2_Calories.std() / (len(df2_Calories)**0.5)


twosample_Z = (abs(x1 - x2) - 0) / (se1**2 + se2**2)**0.5

twosample_Z
```

Output:
```
1.025155364042557
```

<!-- grade_id: cell-c2ae9a30fd101418, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Part Puzzle2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert( abs(round(twosample_Z,5)) == 1.02516), "twosample z-score is not correct"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
```

Output:
```
🎉 All tests passed! 🎉
```

## 3. Submit Your Work

In this notebook:

1. Click **File** -> **Save and Checkpoint** (to save your work)
2. Click **File** -> **Close and Halt** (to exit this notebook)

Follow the instructions on the STAT 107 website to submit your work.