## EC Notebook for Lecture 37: t Tests

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30 am on **Wednesday, Dec. 04** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## 1. One Sample t Test

In this EC notebook, we will use the `spiffy` dataset and `scipy.stats.ttest_1samp` tool to do a t-test. So firstly, let's import the libraries and load the data.

<!-- grade_id: cell-351d38d62038f3f5, points: 0, grade: false, solution: true -->
```python
# load your libraries
df = ...


df
```

<!-- grade_id: cell-bfb8f8c7929ded38, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(len(df) == 8), "The df is incorrect."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
```

Let's try the example we've seen in the lecture,

> A car company claims that their Super Spiffy Sedan averages 31 mpg. You randomly select 8 Super Spiffies from local car dealerships and test their gas mileage under similar conditions.  Assume the gas mileage follows the normal curve.  You get the MPG scores and load them in `df`. Do a t-test in Python to figure out if the mileage is statistically different from 31.

Note that in dataset, there are two columns, `mpg-spiffy` contains sampled mpg of Super Spiffies, and `mpg-another` stands for sampled mpg for another version of car used to do comparison with spiffies.

Create a t test sample and store it in variable `res`.

<!-- grade_id: cell-e944a7df4529ca47, points: 0, grade: false, solution: true -->
```python

res = ...


print("The t-statistic is %.3f and the p-value is %.3f." % res)
```

<!-- grade_id: cell-16cf28eb7c11ad0e, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(abs(res[0]+1.265) < 0.001), "The t-statistic is incorrect."
assert(abs(res[1]-0.246) < 0.001), "The p-value is incorrect."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
```

##  2. Two Sample t Test

Now, let's compare the mpg of spiffy and another type of cars. What is t-statistics and what is 2-tail p-value? *(You can use `from scipy.stats import t` and `t.cdf(0.5, df=10)` to get the p-value when `t=0.5 and df=10`)*

<!-- grade_id: cell-fef161fc8a6e05b2, points: 0, grade: false, solution: true -->
```python
t_statistics = ...
p_value = ...


print("The t-statistic is %.3f and the p-value is %.3f." % (t_statistics, p_value))
```

<!-- grade_id: cell-6f61254e771fe914, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for 3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(abs(t_statistics+0.584) < 0.001), "The t-statistic is incorrect."
assert(abs(p_value-0.569) < 0.001), "The p-value is incorrect."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
```

## 3. Submit Your Work

In this notebook:

1. Click **File** -> **Save and Checkpoint** (to save your work)
2. Click **File** -> **Close and Halt** (to exit this notebook)

Follow the instructions on the STAT 107 website to submit your work.

```python

```