# EC Notebook for Lecture 15: Probability

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30 am on **Wednesday, Oct. 2** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## 2 Control Flow in Python
Remember the coin-tossing we did in the last EC Notebook? If not, that's ok. What you need to do is as easy as tossing a coin! But instead of tossing the coin 10 times, today we want to do it 100 times and even more! Because we already get a secret weapon -- the Loop.

Firstly, remember to import random.
As usual, let's say Head is represented by 1 and Tail is represented by 0. We will toss the coin several times and calculate how many times Head appeared.

### Question 1: The same as last time
Toss the coin 10 times and calculate the probability Head appeared. (Of course you can copy the answer from the last notebook, or try a new way.)

<!-- grade_id: cell-99d959c1fa5c1922, points: 0, grade: false, solution: true -->
```python
import random
toss1 = random.randint(0,1)
toss2 = random.randint(0,1)
toss3 = random.randint(0,1)
toss4 = random.randint(0,1)
toss5 = random.randint(0,1)
toss6 = random.randint(0,1)
toss7 = random.randint(0,1)
toss8 = random.randint(0,1)
toss9 = random.randint(0,1)
toss10 = random.randint(0,1)



# Some basic calculations based on your tosses:
heads_ct = toss1 + toss2 + toss3 + toss4 + toss5 + toss6+ toss7 + toss8 + toss9 + toss10
tails_ct = 10 - heads_ct
probability_head = heads_ct / 10

```

<!-- grade_id: cell-6d95af9e6dbf7763, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(0 <=probability_head <=1), "The probability is always between 0 and 1, isn't it?"
```

### Question 2: Toss more times!
Toss the coin **100** times and calculate the probability Head appeared. 

<!-- grade_id: cell-0f2b9437c5b3e29a, points: 0, grade: false, solution: true -->
```python
# How about store the current count of head in a new variable?
head_count = 0
for i in range(100):
    toss = random.randint(0,1)
    head_count += toss 

probability_head_100 = head_count / 100
probability_head_100
    



```

Output:
```
0.51
```

<!-- grade_id: cell-6f4ebbb95f35fe62, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(0 <=probability_head_100 <=1), "The probability is always between 0 and 1, isn't it?"
```

### Question 2: Toss more coins!
Toss **two** coins **100** times and calculate the probability **TWO Heads** appeared at the same time. 

<!-- grade_id: cell-ea88dcf40cc90ada, points: 0, grade: false, solution: true -->
```python
two_head_count = 0
for i in range(100):
    toss1 = random.randint(0,1)
    toss2 = random.randint(0,1)
    # Hint 1: if both coins are head, toss1 + toss2 == 2
    # Hint 2: we can convert a real number to an integer using `int()`
    # For example: `a = int(5.6)`, then we get an ingeger `a` equal to 5.
    # more reference: https://www.tutorialspoint.com/python/python_numbers.htm
    # Plus, if you know the `if` python syntax, feel free to use it.
    count = toss1 + toss2
    
    if count == 2:
        two_head_count += 1
print(two_head_count)
probability_two_head = two_head_count / 100
probability_two_head
    



```

Output:
```
29
0.29
```

<!-- grade_id: cell-357bac434e050bca, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(0 <=probability_two_head <=0.5), "I don't think the probability is correct..."
```

## Submit Your Work!

Make sure to **Save and Checkpoint** your notebook, exit Jupyter, and submit your work! :)

```python

```