## EC Notebook for Lecture 10:  Functions in Python

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30am on **Friday, Sept. 20** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.


## 1. Importing pandas and reading data
For this extra credit notebook, we will still use the Illinois GPA Distribution dataset -- as we did in lecture -- to do some exercise to get familiar with the function we learnt in the lecture

You feel free to access the "GPA" dataset here: https://github.com/wadefagen/datasets/tree/master/gpa
At first, you need to import the library pandas and reading the dataset.

<!-- grade_id: cell-b94d835c3afb8ee1, points: 0, grade: false, solution: true -->
```python
import pandas as pd 
df = pd.read_csv('gpa.csv')
df
```

Output:
```
Year    Term YearTerm Subject  Number                   Course Title  \
0      2018    Fall  2018-fa     AAS     100   Intro Asian American Studies   
1      2018    Fall  2018-fa     AAS     100   Intro Asian American Studies   
2      2018    Fall  2018-fa     AAS     100   Intro Asian American Studies   
3      2018    Fall  2018-fa     AAS     100   Intro Asian American Studies   
4      2018    Fall  2018-fa     AAS     100   Intro Asian American Studies   
...     ...     ...      ...     ...     ...                            ...   
46540  2010  Summer  2010-su    STAT     410  Statistics and Probability II   
46541  2010  Summer  2010-su    STAT     440    Statistical Data Management   
46542  2010  Summer  2010-su     TAM     212          Introductory Dynamics   
46543  2010  Summer  2010-su     TAM     251   Introductory Solid Mechanics   
46544  2010  Summer  2010-su     TAM     335   Introductory Fluid Mechanics   

       A+   A  A-  B+  ...  B-  C+  C  C-  D+  D  D-  F  W
```

## 2. Review: Group Function
The Student Union at Illinois would like to investiage the grading policy for the university courses and subjects， especailly for "ACCY", "CS", "MATH", "STAT".

### Question 1: 
Using the GPA dataset, you need to calculate the proportion of ("A+","A") in ("A+","A","A-","B+","B","B-") for subjects "ACCY", "CS", "MATH", "STAT", respectively. Hint: use the groupby function.

<!-- grade_id: cell-b1e450db85c95e49, points: 0, grade: false, solution: true -->
```python
df_s = df[ (df['Subject'] == 'ACCY') | (df['Subject'] == 'CS') | (df['Subject'] == 'MATH') | (df['Subject'] == 'STAT') ]
group = df_s.groupby('Subject')
df_group = group.agg("sum").reset_index()
ratio = (df_group['A+'] + df_group['A']) / (df_group['A+'] + df_group['A'] + df_group['A-'] + df_group['B+'] + df_group['B'] + df_group['B-']) 
print(ratio[1])



```

Output:
```
0.4851959027145017
```

<!-- grade_id: cell-fc6499a6e7186f28, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Question 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(ratio[1] == 0.4851959027145017), "The ratio for the ACCY does not appear correct."
```

## 3. Review: Group Function
Besides the raito of ("A+","A") in ("A+","A","A-","B+","B","B-"), the student union also wants to know the difference between the largest number of "A+" and the largest number of "A" in each subject, and what is the largest gap among all subjects. 

### Question 2: 
Using the GPA dataset, you need to find the difference between the largest number of "A+" and the largest number of"A" for each subject firstly, and then based on what you have, to calculate the largest gap among all subjects. 

<!-- grade_id: cell-17f1a17a9291bd92, points: 0, grade: false, solution: true -->
```python

group_all = df.groupby('Subject')
df_gMax = group_all.agg("max").reset_index()
df_gMax["Max Diff"] = df_gMax['A+'] - df_gMax['A']
df_gMax["Max Abs Diff"] = df_gMax['Max Diff'].abs()
gap_largest = df_gMax["Max Abs Diff"].max()
print(gap_largest)



```

Output:
```
481
```

<!-- grade_id: cell-63f5df829092ffa4, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Question 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(gap_largest == 481), "The number does not appear correct."
```

## 4. Review: Group Function
In addition, the student union also wants to know hom many subjects which the instructors give "A" but not "A+". 


### Question 3: 
Using the GPA dataset, we have defined a function to give the label (column "course_type") for those course which the instructors give "A" but not "A+". For your work, you need to decide how many subjects which the instructors give "A" but not "A+". Hint: The question is asking about each subject not each course, so you also need to use groupby function.

<!-- grade_id: cell-656278288cff2886, points: 0, grade: false, solution: true -->
```python
## Just Run the following code. 
def course(x):

   if x["A+"] == 0 and x["A"] > 0 : return 1

   else: return 0
    
df["course_type"] = df.apply(course, axis=1)
## Just Run the above code. 

df_mean = group_all.agg("mean").reset_index()
df_num = df_mean[df_mean['course_type'] == 0]
num = len(df_num)
num
```

Output:
```
13
```

<!-- grade_id: cell-ce3a059c412e2e95, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Question 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(num == 13), "The number of this type of course does not appear correct."
```

```python

```