## EC Notebook for Lecture 27: CLT + Polling

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30 am on **Monday, Nov. 4** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## 1. Function to calculate the number of people needed in the survey

In the class, we saw a formula to calculate the number of people needed for a survey (with known underlying SD) to achieve confidence interval (level alpha) with a certain margin of error. <br>

We will use the following notations: <br>

**margin** : represents the margin of error required. <br>
**SD** : underlying population standard deviation. <br>
**confidence** : the confidence percentage. <br>
**n** : the number of people required

Write a function which takes **margin, SD** and **confidence** as inputs and outputs the number of people (**n**). <br>
Reacall that the formula which we can use is as: <br>

**n=(100 * z_confidence*SD/margin)^2**, <br>

where **z_confidence** = **z_value** corresponding to **confidence** (example: 95% confidence has z_value = 1.96) <br>

Note that the confidence is given in percentage. So, make sure to change it into a fraction while trying to calculate z_value.

<!-- grade_id: cell-3182d2cfb9079fe2, points: 0, grade: false, solution: true -->
```python
from scipy.stats import norm
import math
def n_calc(margin, SD, confidence):
    #first calculate z_value
    z_confidence = norm.ppf((1 - (confidence/100))/2 + (confidence/100))
    #next, calculate and return n
    n= (100*z_confidence*(SD/margin))**2
    return math.ceil(n)



```

<!-- grade_id: cell-621e2757b30ebf38, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(n_calc(2,0.5,95) == 2401), "Yikes! Make sure to double check that z_value is correct."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()

```

Output:
```
🎉 All tests passed! 🎉
```

## 2: The formula in action!

### Part 1:
We will write a simulation to toss a coin `n` times to estimate the probability of getting heads. Note that its population value (true value) is 0.5 (which we don't know!). <br>
Before we write the simulation, suppose we want to estimate the probability with `4%` margin of error with `90%` confidence. First, calculate the number of tosses needed to achieve that. (**SD = 0.5**) <br>
Hint: Use the function which you have defined above!

<!-- grade_id: cell-2ff20b5b160999ef, points: 0, grade: false, solution: true -->
```python
n_toss = n_calc(4,0.5,90)



```

### Part 2:

Now, write the simulation of tossing a coin `n_toss` many times. Calculate the margin of error that you get at `90%` confidence. Is it close to 4? <br>
Consider assigning **Heads = 1 and Tails = 0**.

<!-- grade_id: cell-2df596612092f1ba, points: 0, grade: false, solution: true -->
```python
z_value = norm.ppf(0.95)
import random
import pandas as pd
import math

data = []
for i in range(n_toss):
    toss = random.randint(0,1)
    d = {"toss": toss}
    data.append(d)
df = pd.DataFrame(data)

se = 0.5/math.sqrt(n_toss)
margin = 100*z_value*se    
margin


```

Output:
```
3.998776005175958
```

```python

```