# EC Notebook for Lecture 13: Algorithms

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30 am on **Friday, Sept. 27** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## 1. Loading pandas and importing data

In this extra credit notebook, we will still exploring the information hidden in the "Illini Football" dataset -- as we did in lecture 13 -- to answer some questions.

you can access the "Illini Football" dataset here: 
https://raw.githubusercontent.com/wadefagen/datasets/master/illini-football/illini-football-scores.csv

Next, you need to load the pandas and dataset.

<!-- grade_id: cell-c08171d36379ef97, points: 0, grade: false, solution: true -->
```python
import pandas as pd 

df = pd.read_csv('illini-football-scores.csv')
df
```

Output:
```
Season        Date       Location                Opponent Result  \
0       2018  2018-09-01            vs.              Kent State      W   
1       2018  2018-09-08            vs.        Western Illinois      W   
2       2018  2018-09-15    Chicago, IL           South Florida      L   
3       2018  2018-09-21            vs.              Penn State      L   
4       2018  2018-10-06              @                 Rutgers      W   
5       2018  2018-10-13            vs.                  Purdue      L   
6       2018  2018-10-20              @               Wisconsin      L   
7       2018  2018-10-27              @                Maryland      L   
8       2018  2018-11-03            vs.               Minnesota      W   
9       2018  2018-11-10              @                Nebraska      L   
10      2018  2018-11-17            vs.                    Iowa      L   
11      2018  2018-11-24              @            Northwestern      L   
12      2017  2017-09-02            vs.     
```

## 2 Puzzle Questions from Lecture.
We have load the csv dataset to the python, and let us answer some intersting questions according to the information which the dataset provides for us. 

### Question 1:
How many unique opponents have we played against? 
How many times have we played against each opponent?
What is the Illini win-loss record against Purdue and Chicago, (that is Win - Lost)? 
How many homecoming games have we played?
What is our Win-Loss record among homecoming games, (that is Win - Lost)?

<!-- grade_id: cell-4e3bbaa087fd76f9, points: 0, grade: false, solution: true -->
```python


group = df.groupby('Opponent')
grouped_count = group.agg('count').reset_index()
uniq_opponent = len(grouped_count['Opponent'])

times_each = df['Opponent'].value_counts()


df_purdue = df[df['Opponent'] == 'Purdue']
df_chicago = df[df['Opponent'] == 'Chicago']
result_p = df_purdue['Result'].value_counts()
result_c = df_chicago['Result'].value_counts()

wins_lost_purdue = result_p[0] - result_p[1]
wins_lost_chicago = result_c[0] - result_c[1]

df_hc = df[df['Note'] == 'Homecoming']
num_homecoming = len(df_hc)

result_hc = df_hc['Result'].value_counts()
wins_lost_homecoming = result_hc[1] - result_hc[0]



times_each
```

Output:
```
Northwestern              112
Ohio State                103
Michigan                   95
Purdue                     93
Wisconsin                  85
Iowa                       74
Minnesota                  71
Indiana                    71
Michigan State             46
Chicago                    43
Penn State                 24
Missouri                   24
Washington                 20
Nebraska                   16
Southern California        13
Notre Dame                 12
Syracuse                   12
UCLA                       12
Stanford                   10
California                 10
Butler                      9
Bradley                     8
Army                        8
Pittsburgh                  8
Wabash                      8
Washington State            7
Knox                        7
Physicians & Surgeons       7
Western Michigan            6
Iowa State                  6
                         ... 
Florida A&M                 1
Oklahoma                    1
Youngstown
```

<!-- grade_id: cell-07f8399e796c5d46, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Question 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(uniq_opponent  == 131), "The unique opponent number does not appear correct."
assert(times_each[0] == 2), "The times against each opponent does not appear correct."
assert(wins_lost_chicago == 4), "The wins - lost against Chicago does not appear correct."
assert(num_homecoming == 108), "The number of Homecoming does not appear correct."
assert(wins_lost_homecoming == -18), "The wins - lost at Homecoming does not appear correct."
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-34-30bd10b39227> in <module>
      5 
      6 assert(uniq_opponent  == 131), "The unique opponent number does not appear correct."
----> 7 assert(times_each[0] == 2), "The times against each opponent does not appear correct."
      8 assert(wins_lost_chicago == 4), "The wins - lost against Chicago does not appear correct."
      9 assert(num_homecoming == 108), "The number of Homecoming does not appear correct."

AssertionError: The times against each opponent does not appear correct.
```

## 3 Algorithm
When we analyze the dataset, we are only curious about some certain columns of the dataset, so it is intuitively to subset the whoel dataset by removing some columns or just select some columns. 


### Question 2:
Try to subset the "df" dataset as a new dataset called "dfnew" which dose not include "Season", "Date", "Note".
Now, I believe you have got the dataset "dfnew", let us calculate the mean, median and standard deviation for IlliniScore at home ("vs.")

<!-- grade_id: cell-8a5e50c5c1967623, points: 0, grade: false, solution: true -->
```python
dfnew = df.drop(columns=['Season','Date','Note'])
dfnew_home = dfnew[dfnew['Location'] == 'vs.']
mean = dfnew_home['IlliniScore'].mean()
median = dfnew_home['IlliniScore'].median()
sd = dfnew_home['IlliniScore'].std()


```

<!-- grade_id: cell-92108dd37135594b, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Question 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(round(mean,0)  == 21.0), "The mean of IlliniScore at home does not appear correct."
assert(median == 20), "The median of IlliniScore at home does not appear correct."
assert(round(sd,0) == 16.0), "The standard deviation of IlliniScore at home does not appear correct."
```

```python

```