## EC Notebook for Lecture 32: RMSE and Clustering

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30 am on **Wednesday, Nov. 13** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

### 1. Importing libraries and reading data
For this extra credit notebook, we will still use the Beer dataset to do some exercise. At first, you need to import the libraries and reading the dataset.

<!-- grade_id: cell-3182d2cfb9079fe2, points: 0, grade: false, solution: true -->
```python
#import libraries
import pandas as pd 
df = pd.read_csv('beer.csv')



df.sample(5)
```

Output:
```
Brand             Brewery  PercentAlcohol  Calories  \
75  Winehard's Private Reserve              Miller             4.8       150   
26         Hamm's Golden Draft              Miller             4.7       144   
38          Magnum Malt Liquor              Miller             5.6       157   
79         Winehard's Pale Ale              Miller             4.6       147   
20           Genesee Cream Ale  High Falls Brewing             5.1       162   

    Carbohydrates  
75            9.9  
26           12.1  
38           11.2  
79           13.0  
20           15.0
```

### 2. Creating regression models
After loading the data, let's fit several regression models according to the following description:

1. `models[0]`: Take **PercentAlcohol** as independent variable and **Calories** as dependent variable.
2. `models[1]`: Take **Carbohydrates** as independent variable and **Calories** as dependent variable.
3. `models[2]`: Take **PercentAlcohol** and **Carbohydrates** as independent variables and **Calories** as dependent variable.

<!-- grade_id: cell-060c81a4477630d2, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LinearRegression
model0 = LinearRegression()
models0 = model.fit(df[['PercentAlcohol']],df['Calories'])

model1 = LinearRegression()
models1 = model.fit(df[['Carbohydrates']],df['Calories'])

model2 = LinearRegression()
models2 = model.fit(df[['PercentAlcohol','Carbohydrates']],df['Calories'])



```

### 3. Making Predictions and Calculating RMSE
Now make predictions using the above models, create three extra columns to store the results. After that, calculate the RMSE with the help of the predicted values and store the results in a list `errors`. 

*Think about: which value is the smallest and why?*


<!-- grade_id: cell-8b0f0ef67575550d, points: 0, grade: false, solution: true -->
```python
import math
df['predict1'] = models0.predict(df[['Calories']])
df['error2']=models1.predict(df[['Carbohydrates']])
df['error3']=models2.predict(df[['PercentAlcohol','Carbohydrates']])
error1=df['error1']**2
error1=error1.mean()
error1=error1**.5
error2=df['error2']**2
error2=error2.mean()
error2=error2**.5
error3=df['error3']**2
error3=error3.mean()
error3=error3**.5
errors=[error1,error2,error3]
errors

```

Output:
```
~\Anaconda3\lib\site-packages\sklearn\linear_model\base.py in _decision_function(self, X)
    204         X = check_array(X, accept_sparse=['csr', 'csc', 'coo'])
    205         return safe_sparse_dot(X, self.coef_.T,
--> 206                                dense_output=True) + self.intercept_
    207 
    208     def predict(self, X):

~\Anaconda3\lib\site-packages\sklearn\utils\extmath.py in safe_sparse_dot(a, b, dense_output)
    140         return ret
    141     else:
--> 142         return np.dot(a, b)
    143 
    144 

ValueError: shapes (85,1) and (2,) not aligned: 1 (dim 1) != 2 (dim 0)
```

<!-- grade_id: cell-621e2757b30ebf38, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Part 1,2,3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

ans_coef = [34.961178732626145, 5.399999224018191, 26.633421167621528]
ans_int = [-26.289717456127363, 82.27177327363447, -11.764312556069427]
ans_error = 38.152395545641205
for i in range(3):
    assert(abs(sum(models[i].coef_)-anscoef[i])<0.01), "Make sure you created the correct models"
    assert(abs(models[i].intercept_-ans_int[i])<0.01), "Make sure you created the correct models"
assert(abs(sum(errors)-ans_error)<0.01), "Make sure you calculated the correct RMSE"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()

```

```python

```