## EC Notebook for Lecture 33: While Loops + K Means Clustering

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30 am on **Friday, Nov. 15** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## While Loop

In Wednesday's lecture, we learned how to create and use a **While Loop**. With the **While Loop**, we can execute a set of statements as long as a condition is true.

Now, you'll be able to do a few practice puzzles to get familiar with using the **While Loop**.

## Puzzle 1

Suppose Mike has 300 dollars and two fair six-sided dice. Mike rolls the two dice and calculates the sum of the two numbers appearing. For example: if the first die lands on 1 and the second die lands on 2, the sum would be 3.  

If the sum is larger than 9, then he will **win** 1 dollar, otherwise, he will **lose** 1 dollar. If Mike starts with 300 dollars, calculate how many times he would have to play to only have 20 dollars remaining. 

<!-- grade_id: cell-68f7fc16e910bc77, points: 0, grade: false, solution: true -->
```python
import random
d = 300
times = 0
while d>20:
    roll1 = random.randint(1,6)
    roll2 = random.randint(1,6)
    t = roll1 + roll2
    if t > 9:
        d += 1
    else:
        d -= 1
    times += 1



print(times)



```

Output:
```
392
```

<!-- grade_id: cell-1f7e4ba428d6dc0b, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert( 340 <= times <= 460), "The times of gameplayed is not correct."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
```

Output:
```
🎉 All tests passed! 🎉
```

## K-Means Clustering

In lecture, we learned a new unsupervised learing technique called **k-means** clustering. 

The term **"k-means"** was first used by James MacQueen in 1967, though the idea goes back to Hugo Steinhaus in 1956. The standard algorithm was first proposed by Stuart Lloyd of Bell Labs in 1957 as a technique for pulse-code modulation. 

**k-means** clustering aims to partition $n$ observations into $k$ clusters in which each observation belongs to the cluster with the nearest mean. The **k-means** algorithm starts with a first group of randomly selected centroids, which are used as the beginning points for every cluster, and then performs iterative (repetitive) calculations to optimize the positions of the centroids. Similar to the example we did by hand in class!

Let's get some practice with **k-means** clustering by looking at the classic `beer` dataset we saw in class! 

## Puzzle 2 

In the `beer` dataset, we have the following variables: `PercentAlcohol`, `Calories` and `Carbohydrates`.  These are three characteristics for different type of beers. In the following puzzle, we will practice **k-means** clustering with this dataset.

Use `PercentAlcohol` and `Calories` as the two characteristics to create four clusters using **k-means** clustering.  Next, write Python code to determine how many beers are in **cluster 0**. You will want to create a new variable in the dataframe that says which cluster each beer is in.  Lastly, create a scatter plot using `PercentAlcohol` as your X variable and `Calories` as your Y variable. Make each `cluster` type a different **color** in the scatterplot. 

<!-- grade_id: cell-b7b9a3f22c08ff32, points: 0, grade: false, solution: true -->
```python
import pandas as pd 
df = pd.read_csv('beer.csv')

length_cluster0 = ...


```

<!-- grade_id: cell-3b641b2893ddcf73, points: 0, grade: false, solution: true -->
```python
import matplotlib.pyplot as plt

plot = ...


```

<!-- grade_id: cell-f6065be652ab534b, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert( 20 <= length_cluster0 <= 38), "The times of gameplayed is not correct."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
```

## 3. Submit Your Work

In this notebook:

1. Click **File** -> **Save and Checkpoint** (to save your work)
2. Click **File** -> **Close and Halt** (to exit this notebook)

Follow the instructions on the STAT 107 website to submit your work.