# EC Notebook for Lecture 5: Combining DataFrames

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30am on **Monday, Sept. 9** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## 1. Loading pandas and importing data

For this notebook, we will use the "Course Catalog" dataset that is our dataset of the week.

Here's the link: https://docs.google.com/spreadsheets/d/1ZnpZpipyui_zIIBblmalRYiyEG1glUKinqJtQxVL04o/edit#gid=0

Use the following cell to load the pandas library and read the DataSet as a DataFrame:

<!-- grade_id: cell-ba4c939dda92ddf8, points: 0, grade: false, solution: true -->
```python
import pandas as pd 
df = pd.read_csv("Illinois Course Catalog - Sheet1.csv")
df

```

Output:
```
Year  Term YearTerm Subject  Number  \
0     2019  Fall  2019-fa     AAS     100   
1     2019  Fall  2019-fa     AAS     105   
2     2019  Fall  2019-fa     AAS     120   
3     2019  Fall  2019-fa     AAS     199   
4     2019  Fall  2019-fa     AAS     200   
...    ...   ...      ...     ...     ...   
8584  2019  Fall  2019-fa    ZULU     202   
8585  2019  Fall  2019-fa    ZULU     403   
8586  2019  Fall  2019-fa    ZULU     404   
8587  2019  Fall  2019-fa    ZULU     405   
8588  2019  Fall  2019-fa    ZULU     406   

                                       Name  \
0              Intro Asian American Studies   
1     Introduction to Arab American Studies   
2             Intro to Asian Am Pop Culture   
3                Undergraduate Open Seminar   
4                      U.S. Race and Empire   
...                                     ...   
8584                     Elementary Zulu II   
8585                    Intermediate Zulu I   
8586                   Intermediate Zulu I
```

## 2. Selecting all SUBJ 100 and SUBJ 101

In many subjects, 100 and 101 courses serve as freshman entry-level courses.  Create single DataFrame called `entryCourses` that contains all courses with numbers `100` or `101`.  *There are MANY different ways you can do this, including using `df.merge` if needed.*

<!-- grade_id: cell-ca1503cca7507a80, points: 0, grade: false, solution: true -->
```python
subj_100 = df[ df["Number"] == 100]
subj_101 = df[ df["Number"] == 101]
print(len(subj_100))
print(len(subj_101))
len(subj_100)+len(subj_101)

entryCourses = pd.concat([subj_100,subj_101])
entryCourses







```

Output:
```
48
66
      Year  Term YearTerm Subject  Number                                Name  \
0     2019  Fall  2019-fa     AAS     100        Intro Asian American Studies   
49    2019  Fall  2019-fa     ABE     100  Intro Agric &amp; Biological Engrg   
149   2019  Fall  2019-fa     ACE     100          Agr Cons and Resource Econ   
300   2019  Fall  2019-fa      AE     100      Intro to Aerospace Engineering   
383   2019  Fall  2019-fa    AFRO     100        Intro to African American St   
...    ...   ...      ...     ...     ...                                 ...   
7874  2019  Fall  2019-fa    SPAN     101                Elementary Spanish I   
8147  2019  Fall  2019-fa    THEA     101        Introduction to Theatre Arts   
8360  2019  Fall  2019-fa     UKR     101                   Basic Ukrainian I   
8368  2019  Fall  2019-fa      UP     101       Introduction to City Planning   
8576  2019  Fall  2019-fa    YDSH     101                 Beginning Yiddish I   

                     
```

<!-- grade_id: cell-51c14f7e00a20667, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
assert( len(entryCourses) == 116 ), "`entryCourses` must contain all 116 entry-level courses."
assert( len(entryCourses[entryCourses.Number == 100] ) == 48 ), "`entryCourses` must contain all 48 courses as SUBJ `100`."
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-6-e00cab48aad4> in <module>
      3 # - If this cell runs without any error our output, you PASSED all test cases!
      4 # - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
----> 5 assert( len(entryCourses) == 116 ), "`entryCourses` must contain all 116 entry-level courses."
      6 assert( len(entryCourses[entryCourses.Number == 100] ) == 48 ), "`entryCourses` must contain all 48 courses as SUBJ `100`."

AssertionError: `entryCourses` must contain all 116 entry-level courses.
```

## 3. Creating a `control` and `treatment` group

Using the `entryCourses` you created in Puzzle 2, create a `treatment` and `control` group that each contain 50 random courses to use an experiment you are developing.

*Remember: A course cannot be in both a treatment and control group!*

<!-- grade_id: cell-8ba4be23607128e9, points: 0, grade: false, solution: true -->
```python
# Create the control group:


control = entryCourses.sample(n = 50)
control


```

Output:
```
50
```

<!-- grade_id: cell-e2fbfed754aeba62, points: 0, grade: false, solution: true -->
```python
# Create the treatment group:

pre_treatment = entryCourses[ ~entryCourses.index.isin(control.index)]
treatment = pre_treatment.sample(n = 50)
treatment


```

Output:
```
Year  Term YearTerm Subject  Number  \
528   2019  Fall  2019-fa    AGED     100   
4016  2019  Fall  2019-fa     GER     101   
8123  2019  Fall  2019-fa      TE     100   
6612  2019  Fall  2019-fa      NS     101   
5055  2019  Fall  2019-fa     KIN     101   
6106  2019  Fall  2019-fa     MUS     101   
5544  2019  Fall  2019-fa    MACS     101   
8002  2019  Fall  2019-fa    STAT     100   
577   2019  Fall  2019-fa     AIS     101   
6875  2019  Fall  2019-fa     POL     101   
5213  2019  Fall  2019-fa     LAS     100   
4935  2019  Fall  2019-fa    ITAL     101   
149   2019  Fall  2019-fa     ACE     100   
6786  2019  Fall  2019-fa    PHYS     100   
2194  2019  Fall  2019-fa     CMN     101   
1066  2019  Fall  2019-fa    ARTE     101   
0     2019  Fall  2019-fa     AAS     100   
8576  2019  Fall  2019-fa    YDSH     101   
1533  2019  Fall  2019-fa    BIOE     100   
3887  2019  Fall  2019-fa    GEOG     101   
2633  2019  Fall  2019-fa    DANC     101   
4583  2019  Fall
```

<!-- grade_id: cell-8d7d473f143588c7, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert( len(control) == 50 ), "The control group must contain exactly 50 courses."
assert( len(treatment) == 50 ), "The control group must contain exactly 50 courses."

assert( len(control[control.Number > 101] ) == 0 ), "`control` must contain only entry-level courses."
assert( len(treatment[treatment.Number > 101] ) == 0 ), "`treatment` must contain only entry-level courses."

assert( len(pd.merge(control, treatment, how='inner', on=['Subject', 'Number'])) == 0 ), "Courses in `treatment` and `control` must be different."
```

## Submit Your Work!

Make sure to **Save and Checkpoint** your notebook, exit Jupyter, and submit your work! :)