# EC Notebook for Lecture 25: Sampling and Confidence Intervals

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30 am on **Monday, Oct. 28** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## Confidence Intervals

In the lecture, we learned about calculating standard error for the mean of a sample. Based on the standard error and the sample mean, we saw how to derive confidence intervals for population mean. <br>
We will use the `gpa` dataset in this EC notebook. So first load the `gpa` dataset.

## Question 1: Sampling 

### 1(a) : The whole data (which we don't have access to in practice)
We will look at the number of **A-** in each course. First plot a histogram of number of students who have obtained **A-**. You can adjust the bin size to visualise it properly. Next, calculate the mean of this column (this will be our population mean).

<!-- grade_id: cell-510886ab7ba9b0f6, points: 0, grade: false, solution: true -->
```python
#import pandas
import pandas as pd 
#load the dataset and store it in df
df = pd.read_csv('gpa.csv')
#plot the histogram for "A-"
df['A-'].hist(bins=20)#
# calculate the mean
mean_A_minus = df['A-'].mean()


```

<!-- grade_id: cell-e3af940b3396ec89, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Question 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert( len(df) == 46545), "Please check if you have loaded the dataset correctly!"
assert( round(mean_A_minus*len(df),1) == 343101), "The mean seems incorrect!"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

Output:
```
🎉 All tests passed! 🎉
```

### 1(b) Small sample (which we have access to in practice)
Now, sample out 30 courses and calculate following three quantities for the number of students with "A-": <br>
1) sample mean <br>
2) standard error of the sample mean (sem)<br>
3) 90% confidence interval for the mean

<!-- grade_id: cell-de509ab2d91b2a74, points: 0, grade: false, solution: true -->
```python
# Draw 30 samples
sample = df.sample(n=30)

sample_mean = sample['A-'].mean()
sem = sample['A-'].sem()
# z_value is the value such that probability from -z_value to z_value is 0.9 (90%)
from scipy.stats import norm
from math import sqrt
z_value = norm.ppf(0.95)
#left_ci is the leftmost point of our confidence interval
left_ci = sample_mean - z_value*(sem/sqrt(30))
#right_ci is the rightmost point of our confidence interval
right_ci = sample_mean + z_value*(sem/sqrt(30))
#confidence interval! Does population mean lie in this?
print([left_ci, right_ci])


```

Output:
```
[6.779371863275629, 7.620628136724371]
```

<!-- grade_id: cell-90c8b537b0ca5a8f, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Question 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert( len(sample) == 30), "Only 30 samples are needed!"
assert(norm.cdf(z_value) == 0.95), "The z_value doesn't seem right!"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

Output:
```
🎉 All tests passed! 🎉
```

## Question 2: Larger sample.

Now, we will repeat the previous part with larger sample size. Take 1000 samples and recalculate the quantities as above. Note how the interval size changes!

<!-- grade_id: cell-0ec527703ee58a75, points: 0, grade: false, solution: true -->
```python
# Draw 1000 samples
sample_new = df.sample(n=1000)

sample_new_mean = sample_new['A-'].mean()
sem_new = sample_new['A-'].sem()
# z_value is the value such that probability from -z_value to z_value is 0.9 (90%)
z_value = norm.ppf(0.95)
#left_ci is the leftmost point of our confidence interval
left_ci_new = sample_new_mean - z_value*(sem_new/sqrt(1000))
#right_ci is the rightmost point of our confidence interval
right_ci_new = sample_new_mean + z_value*(sem_new/sqrt(1000))
#confidence interval! Does population mean lie in this?
print([left_ci_new, right_ci_new])


```

Output:
```
[7.877258264034989, 7.91274173596501]
```

<!-- grade_id: cell-6ecf73f839b3f0f3, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Question 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert( len(sample_new) == 1000), "1000 samples are needed!"
assert(norm.cdf(z_value) == 0.95), "The z_value doesn't seem right!"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

Output:
```
🎉 All tests passed! 🎉
```

## Submit Your Work :)

Make sure to **Save and Checkpoint** your notebook, exit Jupyter, and submit your work! :)