# EC Notebook for Lecture 9: Grouping Data in Pandas

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30am on **Wednesday, Sept. 18** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## 1. Importing pandas and reading data
For this notebook, we will use the Illinois GPA Distribution dataset -- as we did in lecture -- to answer a few questions on working with columns.

You can access the "GPA" dataset here: https://github.com/wadefagen/datasets/tree/master/gpa

<!-- grade_id: cell-9aafb896bbc8e800, points: 0, grade: false, solution: true -->
```python
import pandas as pd 
df = pd.read_csv("gpa.csv")
df
```

Output:
```
Year    Term YearTerm Subject  Number                   Course Title  \
0      2018    Fall  2018-fa     AAS     100   Intro Asian American Studies   
1      2018    Fall  2018-fa     AAS     100   Intro Asian American Studies   
2      2018    Fall  2018-fa     AAS     100   Intro Asian American Studies   
3      2018    Fall  2018-fa     AAS     100   Intro Asian American Studies   
4      2018    Fall  2018-fa     AAS     100   Intro Asian American Studies   
...     ...     ...      ...     ...     ...                            ...   
46540  2010  Summer  2010-su    STAT     410  Statistics and Probability II   
46541  2010  Summer  2010-su    STAT     440    Statistical Data Management   
46542  2010  Summer  2010-su     TAM     212          Introductory Dynamics   
46543  2010  Summer  2010-su     TAM     251   Introductory Solid Mechanics   
46544  2010  Summer  2010-su     TAM     335   Introductory Fluid Mechanics   

       A+   A  A-  B+  ...  B-  C+  C  C-  D+  D  D-  F  W
```

## 2. Review: Boolean Logic and Conditionals
The Illinois CS and STAT Students Centers want to calculate the number of students who did an excellent job in the courses of CS and STAT. In this case, they define those who got an "A+" or "A" are excellent. Also, considering the course difficulty, they think those who got an "A-" in advanced-level courses (course level >= 400) should be counted in excellent students. Can you help them to get the number?

### Question 1: 
Using the GPA dataset, calculate the number of students who got "A+" or "A" (in all level CS and STAT courses) or "A-" (only in advanced-level CS and STAT courses) in the DataFrame `df`. 

<!-- grade_id: cell-c0705f665f8b46dd, points: 0, grade: false, solution: true -->
```python
df_stat = df[df["Subject"] == "STAT"]
df_stat_A = df_stat[df_stat["Number"] >= 400]

df_cs = df[df["Subject"] == "CS"]
df_cs_A = df_cs[df_cs["Number"] >= 400]




num_excellent_students = df_stat["A+"].sum() + df_cs["A+"].sum() + df_stat["A"].sum() + df_cs["A"].sum() + df_stat_A["A-"].sum() + df_cs_A["A-"].sum()
print(num_excellent_students)


```

Output:
```
87191
```

<!-- grade_id: cell-98d118325ba78938, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Question 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(num_excellent_students == 87191), "The number does not appear correct."
```

## 3. Review: Adding a New Variable (column)
Thanks to your great job, CS and STAT Students Centers know better about their courses. And they recommend you to other departments' students center, so they're all coming to ask you for a favor. This time they are curious about the percentage of students who got "F" in each course.

### Question 2: 
Using the GPA dataset, calculate the number of students who got "F" of each course and add a new column titled `F rate`. To test your correctness, print the average `F rate` of all the courses.

<!-- grade_id: cell-e0c17bd27629d741, points: 0, grade: false, solution: true -->
```python
df['F rate'] = df['F'] / (df['A+'] + df['A'] + df['A-'] + df['B+'] + df['B'] + df['B-'] + df['C+'] + df['C'] + df['C-'] + df['D+'] + df['D'] + df['D-'] + df['F'] + df['W'])
avg_f_rate = df['F rate'].mean()
avg_f_rate


```

Output:
```
0.01743423227044942
```

<!-- grade_id: cell-576f5fd03e1f7da2, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Question 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(abs(avg_f_rate - 0.017564803426637177) < 0.01), "The rate does not appear correct."
```

```python

```

```python

```